# Foundation of Language Technolofy WS 22/23: Homework 02

Please send your solution as a zip-file containing all *.ipynb files that were provided for this homework.(.ipynb). Include comments in your program code to make it easier readable. 

**Naming template: Group_X_homework_Y.ipynb, Group_X_homework_Y.zip**

Please replace X with your group number and Y with the homework number. Submissions that do not follow these rules will not be considered. 

Please only modify the template in the specified markdown and code cells (e.g. `YOUR CODE / ANSWER / IMPORTS HERE`). 
Some cells are left blank on purpose. Please do not modify these cells, because they are used to autograde your submission.

The deadline for the homework is **Friday, 02/12/2022**. Late submissions will not be accepted.

```python
# Import type hint
from typing import List, Dict
```

## Part A: NLTK (25 points)

We will discuss more NLTK functions in this homework. It is recommended to use the NLTK's conditional frequency distribution (`ConditionalFreqDist`) and list comprehension . However, you can implement your solutions without these built-in functions.

#### A.1 Sentiment in opnion lexicon (4 points)
The NLTK’s “`opinion_lexicon`” corpus contains roughly 7,000 English words that are classified as having either
positive or negative sentiment. The positive and negative words of the corpus can be accessed via the functions
`opinion_lexicon.positive()` and `opinion_lexicon.negative()`, respectively.

```python
import nltk
## Download 'opinion_lexicon' if you haven't already
# nltk.download('opinion_lexicon')
from nltk.corpus import opinion_lexicon
```

```python
# Try it yourself
# This will returns a lexicon of positive words/tokens
# It will only print the first 5 words, and "..." indicates more words in the lexicon
opinion_lexicon.positive()
```

```python
# Same for negative words
opinion_lexicon.negative()
```

#### 1) Count with conditions (2 points)

Your task: Use the whole `opinion_lexicon` corpus to count number of negative and positive words starting with `'dis'`

Hint: You can go through each word in the lexicon using `for ... in ...` as in a list

<!-- grade_id: A-1-1-code, points: 0, grade: false, solution: true -->
```python
def word_start_with_dis() -> (int, int):
    """
    Return the number of negative and positive words starting with 'dis', respectively in opinion_lexicon corpus
    """
    positive = 0
    negative = 0
    ### Count positive words start with "dis"
    # YOUR CODE HERE
    raise NotImplementedError()
    
    ### Count negative words start with "dis"
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return negative, positive
```

<!-- grade_id: A-1-1-public, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 1 point
negative, positive = word_start_with_dis()
assert positive > 0
assert negative > 0
```

<!-- grade_id: A-1-1-private, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 1 point
```

#### 2) Conditional probability (2 points)

Your task: ***Calculate the `probability` of a word having a `negative sentiment`, given that it `starts with "dis"` over the whole lexicon***

To estimate the probability, we need to divide the number of all `negative words` starting with `"dis"` over the sum of `all words` (`positive and negative`) starting with `"dis"`

$$P(\text{negative | starts with "dis"}) = \frac{\text{# negative words starting with "dis"}}{\text{# positive & negative words starting with "dis"}}$$

<!-- grade_id: A-1-2-code, points: 0, grade: false, solution: true -->
```python
def cond_probability(negative: int, positive: int) -> float:
    """
    Calculate the probability of a word having a negative sentiment 
    
    Parameters
    ------
    negative: int
        Number of negative words
    positive: int
        Number of positive words
        
    return: float
        Probability of having a negative sentiment
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: A-1-2-pulic, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 1 point
result = cond_probability(negative, positive)
assert 0 < result < 1, "A probability ranges from 0 to 1"
assert isinstance(result, float), "A probability is a float number"
```

<!-- grade_id: A-1-2-private, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 1 point
```

# A.2 Text Generation (11 points)
In Lecture 4, we show how to generate a sentence using $n$-gram language models built from `english-kjv.txt` corpus.

```python
import nltk
import random
```

```python
# Load words from `english-kjv.txt`
text = nltk.corpus.genesis.words("english-kjv.txt")
# Get bigrams from the text
bigrams_obj = nltk.bigrams(text)
# Use each word as the given condition for the following word 
# and create conditional frequency distribution 
cfd = nltk.ConditionalFreqDist(bigrams_obj)


# Function from Lecture 4
def generate_sentence(
    bigram_model: nltk.ConditionalFreqDist,
    start_word: str, 
    number_of_words=20
) -> str:
    """
    Generate sentence from the given prompt and bigram model  
    
    Parameters
    ------
    bigram_model: ConditionalFreqDist
        Bigram model in form of conditional frequency distribution 
    start_word: str
        Prompt to generate sentence
    number_of_words: int
        Length of the generated sequence
    
    return: string
        A generated text with given bigram language model, start prompt and number of generated words
    """
    prev_word = start_word
    words = [prev_word]
    for i in range(number_of_words):
        # select the next word with the highest probability
        word = bigram_model[prev_word].max()
        prev_word = word
        words.append(word)
    return " ".join(words)
```

```python
# Let's try to generate a sentence from the word `living`
generate_sentence(cfd, "living")
```

#### 1) Build bigram model from a corpus (2 points)

In this task, we are going to use the `austen-persuasion.txt` from `gutenberg` corpus.

Your task: **Load words from the corpus and build a bigram model from it using `nltk.ConditionalFreqDist`**

```python
## Download 'gutenberg' corpus if you haven't already
## nltk.download('gutenberg')
from nltk.corpus import gutenberg
```

<!-- grade_id: A-2-1-code, points: 0, grade: false, solution: true -->
```python
def build_bigram_models_from_austen_persuasion():
    """
    Build a bigram model from words in austen-persuasion.txt
    
    return: ConditionalFreqDist
        Bigram model with words from austen-persuasion.txt
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return bigram_model

bigram_model = build_bigram_models_from_austen_persuasion()
```

```python
# Let's try to generate a sentence with the new bigram model
start_word = "We"
generated_text_w_highest_prob = generate_sentence(bigram_model, start_word)
print("Generated sentence by selecting highest probability: \n{}".format(
    generated_text_w_highest_prob))
```

<!-- grade_id: A-2-1-public, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 1 point
start_word = "We"
generated_text_w_highest_prob = generate_sentence(bigram_model, start_word)
expected_gen_text = "We are not be a very much to be a very much to be a very much to be a very"
assert generated_text_w_highest_prob == expected_gen_text
```

<!-- grade_id: A-2-1-private, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 1 point
```

#### 2) Generate sentence from most common words (9 points)

Your task: **generate a sentence of `20 words` from a given `start_word` by randomly selecting the next word from the `5 most common words` that may `occur next`, using the bigram model built on `austen-persuasion.txt`**

- Words for our bi-grams model are taken from `austen-persuasion.txt` text in `Gutenberg` corpus. 
- The generated text should have exactly `20` words/tokens.

**Hint:** 
- `ConditionalFreqDist.most_common(n: int)` takes a number `n` as input and returns a list of `n` most common words and their counts \[(word_1, n_1), (word_2, n_2), ...\]
- `random.choice` takes input as a list of tokens (top 5 possible tokens given the previous token) and randomly selects one. 

<!-- grade_id: A-2-2-code, points: 0, grade: false, solution: true -->
```python
def generate_sentence_from_most_common_words(
        bigram_model: nltk.ConditionalFreqDist,
        start_word: str,
        number_of_words=20,
        number_of_most_common_words=5
) -> str:
    """
    Generate a sequence of tokens with given start token/word using ConditionalFreqDist()
    
    Parameters
    -------
    start_word: string
        a word/token
        
    return: string
        a generated text from exactly 20 tokens 
    """    
    # Initialize the sentence with the start word
    prev_word = start_word
    list_word = [prev_word]
    
    # Find top 5 words that can possibly follow the current word
    # and randomly select one 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # returns the generated sentence 
    return " ".join(list_word)
```

```python
# Try to print the generated sentence here
start_word = "We"
generated_text_w_most_common = generate_sentence_from_most_common_words(bigram_model, start_word)
print("Generated sentence by randomly selecting most common next words: \n{}".format(
    generated_text_w_most_common))
```

<!-- grade_id: A-2-2-public-1, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# + 1 start_word
# 1 point
assert len(nltk.tokenize.word_tokenize(generated_text_w_most_common)) == 20 + 1, \
    "Check whether the function generates 20 words after the start_word"
```

<!-- grade_id: A-2-2-public-2, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 1 point
assert generated_text_w_highest_prob != generated_text_w_most_common, \
    "Two generate functions should return different sentences"
```

<!-- grade_id: A-2-2-private-1, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
```

<!-- grade_id: A-2-2-private-2, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# We set the random seed to 20 for reproducibility
random.seed(20)
start_word = "We"
```

<!-- grade_id: A-2-2-private-3, points: 3, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# This will test whether the solution use 
# the input values of `number_of_words`
# and `number_of_most_common_words`
```

### A.3 Regular Expression (10 points)

Create a function `trigrams_regex(tokens)`, which finds from a given text all sequences of 3-grams that 
1. have an article, 
2. a potential adjective ending with -ing 
3. and a following word by using **regular expressions only**. 
Make sure that the third element of the 3-gram is a word (not punctuation). 

The function should take a list of tokens as an argument and return a list of all found n-grams. 

Try the function on `austen-persuasion.txt` text from the Gutenberg corpus.

```python
import nltk
## Download 'gutenberg' corpus if you haven't already
## nltk.download('gutenberg')
from nltk.corpus import gutenberg
```

`nltk.TokenSearcher` is class that makes it easier to use regular expressions to search over tokenized strings. The tokenized string is converted to a string where tokens are marked with angle brackets – e.g., '<the><window><is><still><open>'. The regular expression passed to the `findall()` method is modified to treat angle brackets as non-capturing parentheses, in addition to matching the token boundaries; and to have '.' not match the angle brackets. Try to run examples below:

```python
from nltk.text import TokenSearcher
from nltk.book import text1, text5, text9
print(text5.findall("<.*><.*><bro>"))
print(text1.findall("<a>(<.*>)<man>"))
print(text9.findall("<th.*>{3,}"))
```

```python
token_searcher = nltk.text.TokenSearcher(gutenberg.words('austen-persuasion.txt'))
found_tokens = token_searcher.findall(r"<\d+><[A-Za-z]+><[A-Za-z]+><[A-Za-z]+>")
found_tokens
```

`token_searcher.findall(r"<\d+><[A-Za-z]+><[A-Za-z]+><[A-Za-z]+>")` returns a list of list where the innere list consists of 4 matched tokens as follows:

- 1st token `<d+>`: matches a digit (equivalent to `[0-9]`) and matches the previous digit at least one to unlimited times
- 2nd, 3rd, 4th token `<[A-Za-z]+>`: matches a single character in the range between `[A-Z]` or `[a-z]` and matches the previous character at least one to unlimited times

With this pattern, we can search for the chapter number and exactly the 3 words that follow it in `'austen-persuasion.txt'`. Since chapter 2 does not match, this means that the first 3 words of the chapter do not belong to the Latin alphabet, but with a punctuation or a character with diacritic for example.

<!-- grade_id: A-3-code, points: 0, grade: false, solution: true -->
```python
def trigrams_regex(tokens: List[str]) -> List[List[str]]:
    """
    Given a given list of tokens, find all 3-grams from the list that 
    1. starts with an article 
    2. followed by a potential adjective ending with -ing
    3. and any following word
    
    Parameters
    -------
    tokens: list of tokens/word
        e.g., from austen-persuasion.txt
    return: list of list of sequence of 3-grams
        e.g., [[a, ding, word], [an, sing, token]]
    """
    bracketed = nltk.text.TokenSearcher(tokens)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: A-3-public, points: 4, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 2 points
trigrams = trigrams_regex(gutenberg.words('austen-persuasion.txt'))
assert trigrams[0] == ['the', 'beginning', 'and']
assert sum([len(t) for t in trigrams]) / len(trigrams) == 3, "We are looking for trigrams (3-grams)"
```

<!-- grade_id: A-3-private, points: 6, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 4 points
```

---

## Part B: spaCy (10 point)


### B.1 Explore real dataset - 10 Points
In this exercise, you will work with a real dataset of public english-language tweets for the keyword 'lockdown'.

```python
from collections import defaultdict
import spacy
## download the language model if you haven't already
# spacy.cli.download("en_core_web_sm")
nlp = spacy.load('en_core_web_sm')
```

To get started, you will have to read the dataset from the provided `tweets.txt` file. Each line in this file represents a single tweet. You will need to open and read the file before starting the other subtasks. 

#### 1) Tokenize each tweet in the dataset. Use spaCy to solve this task. (2 points)

<!-- grade_id: B-1-1-code, points: 0, grade: false, solution: true -->
```python
def read_data(data_name: str) -> spacy.tokens.doc.Doc:
    """
    Read the file name and tokenize each tweet
    
    Parameters
    -------
    data_name: string
        name of the file
    return: spacy.tokens.doc.Doc
        object from spaCy that each token can be iterated
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# DO NOT MODIFY THIS TEST CELL
doc = read_data('tweets.txt')
```

<!-- grade_id: B-1-1-public, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 1 point
assert doc[0].text == 'Nothing'
assert isinstance(doc, spacy.tokens.doc.Doc), "read_data should return `spacy.tokens.doc.Doc` type"
```

<!-- grade_id: B-1-1-private, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 1 point
```

#### 2) Implement the function `occurence_lowercase`.  (8 points)

The function should calculate the (absolute) number of occurrences of each token that is in lowercase. 
Apply the function to our dataset of tweets and return the result (i.e. 'token: occurence') in descending order. Use spaCy to identify lowercased tokens. (7 points)

**Hint**: Do not lowercase all tokens, instead identify and count all already lowercased tokens. 

<!-- grade_id: B-1-2-code, points: 0, grade: false, solution: true -->
```python
def occurence_lowercase(data) -> Dict[str, int]:
    """
    Counts occurences of all lowercased tokens.
    
    Parameters
    -------
    data: array-like object
        containing tokenized tweets from previous subtask
    return: dict-like object         
        object with tokens and their counts
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: B-1-2-public, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# 2 points
occ = occurence_lowercase(doc)
assert occ['the'] == 1111
assert occ['but'] == 145
```

<!-- grade_id: B-1-2-private, points: 6, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS TEST CELL
# Private tests here
# 6 points
```